```python
%matplotlib inline
```

```python
from nose.tools import *
# Write your imports in the cell below
```

<!-- grade_id: cell-3a5c071a139ade21, points: 0, grade: false, solution: true -->
```python
import numpy as np
import pandas as pd
import matplotlib.pyplot as plt
from sklearn.model_selection import train_test_split, GridSearchCV
from sklearn.linear_model import LogisticRegression
from sklearn.metrics import f1_score, make_scorer
```

```python
np.random.seed(1234)
```

# Model Training and Improvement Lab
## Comparing and selecting models

### 1. Read the data (1 point)
Like in the previous lab, you need to read the Portuguese bank dataset [here](https://archive.ics.uci.edu/ml/machine-learning-databases/00222/). It has been provided for you in the `data` folder.

Read the dataset using `pandas` (you can use the library with the alias `pd`). Save it in the `bank_data` variable.

<!-- grade_id: cell-39fe3843121d7b36, points: 0, grade: false, solution: true -->
```python
bank_data = pd.read_csv("data/bank.csv", sep=';')
```

```python
bank_data.head()
```

Output:
```
age          job  marital  education default  balance housing loan  \
0   30   unemployed  married    primary      no     1787      no   no   
1   33     services  married  secondary      no     4789     yes  yes   
2   35   management   single   tertiary      no     1350     yes   no   
3   30   management  married   tertiary      no     1476     yes  yes   
4   59  blue-collar  married  secondary      no        0     yes   no   

    contact  day month  duration  campaign  pdays  previous poutcome   y  
0  cellular   19   oct        79         1     -1         0  unknown  no  
1  cellular   11   may       220         1    339         4  failure  no  
2  cellular   16   apr       185         1    330         1  failure  no  
3   unknown    3   jun       199         4     -1         0  unknown  no  
4   unknown    5   may       226         1     -1         0  unknown  no
```

```python
bank_data.dtypes
```

Output:
```
age           int64
job          object
marital      object
education    object
default      object
balance       int64
housing      object
loan         object
contact      object
day           int64
month        object
duration      int64
campaign      int64
pdays         int64
previous      int64
poutcome     object
y            object
dtype: object
```

```python
bank_data.describe()
```

Output:
```
age       balance          day     duration     campaign  \
count  4521.000000   4521.000000  4521.000000  4521.000000  4521.000000   
mean     41.170095   1422.657819    15.915284   263.961292     2.793630   
std      10.576211   3009.638142     8.247667   259.856633     3.109807   
min      19.000000  -3313.000000     1.000000     4.000000     1.000000   
25%      33.000000     69.000000     9.000000   104.000000     1.000000   
50%      39.000000    444.000000    16.000000   185.000000     2.000000   
75%      49.000000   1480.000000    21.000000   329.000000     3.000000   
max      87.000000  71188.000000    31.000000  3025.000000    50.000000   

             pdays     previous  
count  4521.000000  4521.000000  
mean     39.766645     0.542579  
std     100.121124     1.693562  
min      -1.000000     0.000000  
25%      -1.000000     0.000000  
50%      -1.000000     0.000000  
75%      -1.000000     0.000000  
max     871.000000    25.000000
```

```python
bank_data.groupby("y").size()/len(bank_data.y)*100
```

Output:
```
y
no     88.476001
yes    11.523999
dtype: float64
```

<!-- grade_id: cell-b183f83b98f16a29, points: 1, grade: true, solution: false -->
```python
# From now on, all test cells might contain hidden tests. If you follow the instructions correctly, 
# your solution will be graded with maximum points
assert_is_not_none(bank_data)
```

### 2. Preprocess the data (1 point)
Separate explanatory features from labels. Save all features (16 columns total) in the variable `bank_features`. Save the labels (corresponding to the `y` column) in the `bank_labels` variable. Rewrite the labels to be `0` and `1` instead of `no` and `yes`: `bank_labels` should be a numeric column.

<!-- grade_id: cell-fb1c7a38947afc06, points: 0, grade: false, solution: true -->
```python
bank_features, bank_labels = None, None
bank_features = bank_data.drop("y", axis = 1)
```

```python
bank_features.head()
```

Output:
```
age          job  marital  education default  balance housing loan  \
0   30   unemployed  married    primary      no     1787      no   no   
1   33     services  married  secondary      no     4789     yes  yes   
2   35   management   single   tertiary      no     1350     yes   no   
3   30   management  married   tertiary      no     1476     yes  yes   
4   59  blue-collar  married  secondary      no        0     yes   no   

    contact  day month  duration  campaign  pdays  previous poutcome  
0  cellular   19   oct        79         1     -1         0  unknown  
1  cellular   11   may       220         1    339         4  failure  
2  cellular   16   apr       185         1    330         1  failure  
3   unknown    3   jun       199         4     -1         0  unknown  
4   unknown    5   may       226         1     -1         0  unknown
```

```python
bank_features.shape
```

Output:
```
(4521, 16)
```

```python
bank_labels = bank_data.y
```

```python
bank_labels.shape
```

Output:
```
(4521,)
```

<!-- grade_id: cell-30a215c478663fb9, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(bank_features)
assert_is_not_none(bank_labels)
```

### 3. Get indicator variables (1 point)
Get indicator (dummy) variables for all categorical columns in `bank_features`. Overwrite the `bank_features` variable to store the new data.

<!-- grade_id: cell-8724b6857c2eb7fc, points: 0, grade: false, solution: true -->
```python
bank_features = pd.get_dummies(bank_features)
```

```python
bank_features.info()
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
RangeIndex: 4521 entries, 0 to 4520
Data columns (total 51 columns):
age                    4521 non-null int64
balance                4521 non-null int64
day                    4521 non-null int64
duration               4521 non-null int64
campaign               4521 non-null int64
pdays                  4521 non-null int64
previous               4521 non-null int64
job_admin.             4521 non-null uint8
job_blue-collar        4521 non-null uint8
job_entrepreneur       4521 non-null uint8
job_housemaid          4521 non-null uint8
job_management         4521 non-null uint8
job_retired            4521 non-null uint8
job_self-employed      4521 non-null uint8
job_services           4521 non-null uint8
job_student            4521 non-null uint8
job_technician         4521 non-null uint8
job_unemployed         4521 non-null uint8
job_unknown            4521 non-null uint8
marital_divorced       4521 non-null uint8
marital_married        4521 non-n
```

<!-- grade_id: cell-c2b2c0bd2d30ea5e, points: 1, grade: true, solution: false -->
```python
assert_equal(bank_features.shape, (4521, 51))
```

### 4. Split the data (1 point)
Split the data into training and testing set, with 70% of the data for training. Because the output labels are not equaly distributed, use stratification based on the `bank_labels`.

<!-- grade_id: cell-f46c32bb67654092, points: 0, grade: false, solution: true -->
```python
bank_features_train, bank_labels_train = None, None
bank_features_test, bank_labels_test = None, None
bank_features_train, bank_features_test, bank_labels_train, bank_labels_test = train_test_split(bank_features, bank_labels, train_size = 0.7, stratify = bank_labels)
```

Output:
```
C:\ProgramData\Anaconda3\lib\site-packages\sklearn\model_selection\_split.py:2179: FutureWarning: From version 0.21, test_size will always complement train_size unless both are specified.
  FutureWarning)
```

```python
bank_features_train.shape
```

Output:
```
(3164, 51)
```

```python
bank_labels_train.shape
```

Output:
```
(3164,)
```

```python
bank_features_test.shape
```

Output:
```
(1357, 51)
```

```python
bank_labels_test.shape
```

Output:
```
(1357,)
```

<!-- grade_id: cell-b14f41fb226c0f37, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(bank_features_train)
assert_is_not_none(bank_labels_train)
assert_is_not_none(bank_features_test)
assert_is_not_none(bank_labels_test)
```

### 5. Train a baseline algorithm (1 point)
Train a logistic regression using the training data. Use 1 000 000 (`1e6`) as the value of C. Score it using the testing data. Save the score in the `baseline_score` variable. You should see a fairly high score.

<!-- grade_id: cell-2fd7633e6266fcba, points: 0, grade: false, solution: true -->
```python
model = None
baseline_score = None

model = LogisticRegression(C = 1e6)
```

```python
model.fit(bank_features_train, bank_labels_train)
baseline_score = model.score(bank_features_test, bank_labels_test)
baseline_score
```

Output:
```
C:\ProgramData\Anaconda3\lib\site-packages\sklearn\linear_model\logistic.py:433: FutureWarning: Default solver will be changed to 'lbfgs' in 0.22. Specify a solver to silence this warning.
  FutureWarning)
0.8975681650700074
```

<!-- grade_id: cell-a3af8c95790cccc4, points: 1, grade: true, solution: false -->
```python
assert_is_not_none(model)
assert_greater(baseline_score, 0.7)
```

### 6. Select a better score (2 points)
As you alrady saw, the positive examples are very few. If you aren't convinced, just check the counts.

We know that the default scoring (accuracy) isn't correct in this case. Better measures would be precision and recall. However, we only want one number. Evaluate the algorithm once again, using a standard scoring method which combines precision and recall. Overwrite the `baseline_score` variable.

Don't forget to score the model on the testing data only.

<!-- grade_id: cell-9b44f24da39ba641, points: 0, grade: false, solution: true -->
```python
train_predictions = model.predict(bank_features_train)
test_predictions = model.predict(bank_features_test)
```

```python
train_predictions
```

Output:
```
array(['no', 'yes', 'no', ..., 'no', 'no', 'no'], dtype=object)
```

```python
baseline_score = f1_score(bank_labels_test, test_predictions, pos_label = 'yes')
print(baseline_score)
```

Output:
```
0.4034334763948498
```

<!-- grade_id: cell-9f4aca08f036d9de, points: 2, grade: true, solution: false -->
```python
assert_less(baseline_score, 0.7)
```

### 7. Tune your model (2 points)
Fine-tune the `C` and `max_iter` parameters.

Use full grid search with the following values:
* `C`: 0.0001, 0.01, 0.1, 1, 10, 100, 10000
* `max_iter`: 50, 100, 300, 1000
* `fit_itercept`: True, False

Save the grid search result in the `grid_search` variable. Don't forget to use the better scoring model that you obtained in the previous task.

<!-- grade_id: cell-88cb1341dfd72476, points: 0, grade: false, solution: true -->
```python
grid_search = None

params = {
    'C':[0.0001, 0.01, 0.1, 1, 10, 100, 10000],
    'max_iter':[50, 100, 300, 1000],
    'fit_intercept':[True, False]
}
grid_search = GridSearchCV(LogisticRegression(), params, make_scorer(f1_score, pos_label = 'yes'))
grid_search.fit(bank_features, bank_labels)
```

Output:
```
C:\ProgramData\Anaconda3\lib\site-packages\sklearn\model_selection\_split.py:2053: FutureWarning: You should specify a value for 'cv' instead of relying on the default value. The default value will change from 3 to 5 in version 0.22.
  warnings.warn(CV_WARNING, FutureWarning)
C:\ProgramData\Anaconda3\lib\site-packages\sklearn\linear_model\logistic.py:433: FutureWarning: Default solver will be changed to 'lbfgs' in 0.22. Specify a solver to silence this warning.
  FutureWarning)
C:\ProgramData\Anaconda3\lib\site-packages\sklearn\linear_model\logistic.py:433: FutureWarning: Default solver will be changed to 'lbfgs' in 0.22. Specify a solver to silence this warning.
  FutureWarning)
C:\ProgramData\Anaconda3\lib\site-packages\sklearn\linear_model\logistic.py:433: FutureWarning: Default solver will be changed to 'lbfgs' in 0.22. Specify a solver to silence this warning.
  FutureWarning)
C:\ProgramData\Anaconda3\lib\site-packages\sklearn\linear_model\logistic.py:433: FutureWarning: Default s
```

```python
grid_search.best_estimator_
```

Output:
```
LogisticRegression(C=100, class_weight=None, dual=False, fit_intercept=False,
          intercept_scaling=1, max_iter=50, multi_class='warn',
          n_jobs=None, penalty='l2', random_state=None, solver='warn',
          tol=0.0001, verbose=0, warm_start=False)
```

```python
grid_search.best_params_
```

Output:
```
{'C': 100, 'fit_intercept': False, 'max_iter': 50}
```

<!-- grade_id: cell-0e7311dbe13cf894, points: 2, grade: true, solution: false -->
```python
assert_is_not_none(grid_search)
assert_is_not_none(grid_search.best_estimator_)
```

### 8. Compare scores (1 point)
Use the best estimator from your grid search. Score it using the function from problem 6. Save your answer in `tuned_score`.

<!-- grade_id: cell-98568f584f6c0778, points: 0, grade: false, solution: true -->
```python
tuned_score = None
model = LogisticRegression(C = 100, fit_intercept = False, max_iter = 50)
model.fit(bank_features_train, bank_labels_train)
```

Output:
```
C:\ProgramData\Anaconda3\lib\site-packages\sklearn\linear_model\logistic.py:433: FutureWarning: Default solver will be changed to 'lbfgs' in 0.22. Specify a solver to silence this warning.
  FutureWarning)
LogisticRegression(C=100, class_weight=None, dual=False, fit_intercept=False,
          intercept_scaling=1, max_iter=50, multi_class='warn',
          n_jobs=None, penalty='l2', random_state=None, solver='warn',
          tol=0.0001, verbose=0, warm_start=False)
```

```python
test_predictions = model.predict(bank_features_test)
tuned_score = f1_score(bank_labels_test, test_predictions, pos_label = 'yes')
```

<!-- grade_id: cell-545b4d41fef4eeef, points: 1, grade: true, solution: false -->
```python
print(tuned_score)
```

Output:
```
0.3913043478260869
```

```python
print(baseline_score - tuned_score)
```

Output:
```
0.01212912856876286
```

Hmmmm, it seems we have not obtained a better algorithm, even the opposite (the difference is marginal and depends on the random initialization of the cross-validation datasets).

We can, of course, do a lot more things to improve our model's performance, such as normalizing the data, feature selection and feature engineering, trying out different aspects, e.g. polynomial terms, RANSAC; even boosting (we'll talk about this later). However, we'll stop at this point.

What can we conclude? It seems that this is close to the best performance we can get out of this algorithm, given these data points.

We can try improving (cleaning) our dataset, selecting features, etc. but we most likely need a better algorithm. In the next labs, we're going to explore that.